<problem display_name="Property 3" linenum="127" tex_filename="./block4/innerproducts.tex" showanswer="finished" attempts="2" rerandomize="per_student">
  <text>
    <p>
Let's consider the dot product of a vector with itself. If [mathjaxinline]v \in \mathbb R^3[/mathjaxinline], what is the smallest possible value of [mathjaxinline]\langle v, v\rangle[/mathjaxinline]? </p>
    <p>
      <numericalresponse answer="0">
        <textline>
          <responseparam type="tolerance" default="0.00001"/>
        </textline>
      </numericalresponse>
    </p>
    <p>
For what vector [mathjaxinline]v[/mathjaxinline] is that value of [mathjaxinline]\langle v, v\rangle[/mathjaxinline] attained? </p>
    <p>
      <div class="hideshowbox">
        <h4 onclick="hideshow(this);" style="margin: 0px">How to Enter a Column Vector in R^n<span class="icon-caret-down toggleimage"/></h4>
        <div class="hideshowcontent">
          <p>
To enter a vector in [mathjaxinline]\mathbb R^ n[/mathjaxinline] such as [mathjaxinline]\left[\begin{array}{c} 1 \\ 2 \\ 3 \end{array} \right][/mathjaxinline], you can either: </p>
          <ul class="itemize">
            <li>
              <p>
enter it as you would a [mathjaxinline]3\times 1[/mathjaxinline] matrix: [[1],[2],[3]] </p>
            </li>
            <li>
              <p>
or, you can enter it as &lt;1,2,3&gt; </p>
            </li>
          </ul>
          <p>
Use decimals only. </p>
        </div>
        <p class="hideshowbottom" onclick="hideshow(this);" style="margin: 0px">
          <a href="javascript: {return false;}">Show</a>
        </p>
      </div>
    </p>
    <script linenum="2" filename="./vectorentry2.tex" type="text/python" system_path="python_lib">
def VectorEntry(expect, ans):
	import ast
	import numpy as np 
  	atol = 0.01
	list_expect = ast.literal_eval(expect)
	vec_expect = np.matrix(list_expect)
  	ret = {"ok":False}
	try:
  		# input format [[1],[2],[3]]
		list_ans = ast.literal_eval(ans)
		vec_ans = np.matrix(list_ans)
  		if vec_ans.shape != vec_expect.shape:
  			ret['msg'] = 'Wrong shape of vector!'
  		elif np.allclose(vec_ans, vec_expect,atol,1e-08):
  			ret['ok'] = True
  		else:
  		# More error message. Will improve this part
  			ret['msg'] = 'something is wrong'   			
	#except SyntaxError:
		#ret['msg'] = 'Wrong input format'
	except SyntaxError:
  		# input format &amp;lt;1,2,3&amp;gt;
		list_ans = ans.replace('&amp;lt;', '[').replace('&amp;gt;', ']')
		list_ans = ast.literal_eval(list_ans)
		vec_ans = np.transpose(np.matrix(list_ans))
  		if vec_ans.shape != vec_expect.shape:
  			ret['msg'] = 'Wrong shape of vector!'
  		elif np.allclose(vec_ans, vec_expect,0.01,1e-08):
  			ret['ok'] = True
  		else:
    		# More error message. Will improve this part
  			ret['msg'] = 'something is wrong' 
  	except:
  		ret['msg'] = 'Wrong input format'
  	return ret 
</script>
    <p>
      <customresponse cfn="VectorEntry" expect="[[0],[0],[0]]">
        <textline correct_answer="[[0],[0],[0]]"/>
      </customresponse>
    </p>
    <p>
      <solution>
        <font color="blue">Answer: </font>
        <font color="blue">
          <p>
 If [mathjaxinline]v = \left[ \begin{array}{c} a_1 \\ a_2 \\ a_3 \end{array} \right][/mathjaxinline], then </p>
          <table id="a0000000141" class="equation" width="100%" cellspacing="0" cellpadding="7" style="table-layout:auto">
            <tr>
              <td class="equation" style="width:80%;vertical-align:middle;text-align:center;border-style:hidden">[mathjax]v\dotprod v = a_1^2 + a_2^2 + a_3^2.[/mathjax]</td>
              <td class="eqnnum" style="width:20%;vertical-align:middle;text-align:left;border-style:hidden"/>
            </tr>
          </table>
          <p>
 This is always non-negative, and the only way it can equal zero is if all [mathjaxinline]a_ i[/mathjaxinline] are zero, i.e., if [mathjaxinline]v= \mathbf{0}[/mathjaxinline].  </p>
        </font>
      </solution>
    </p>
  </text>
  <SCRIPT src="/static/latex2edx.js" type="text/javascript"/>
  <LINK href="/static/latex2edx.css" rel="stylesheet" type="text/css"/>
</problem>
