<problem display_name="Closest Vector" linenum="34" tex_filename="./block4/projections.tex" showanswer="finished" attempts="3" rerandomize="per_student">
  <text>
    <p>
Let [mathjaxinline]W[/mathjaxinline] be the subspace of [mathjaxinline]\mathbb R^3[/mathjaxinline] given by </p>
    <table id="a0000000193" class="equation" width="100%" cellspacing="0" cellpadding="7" style="table-layout:auto">
      <tr>
        <td class="equation" style="width:80%;vertical-align:middle;text-align:center;border-style:hidden">[mathjax]W = \left\{  \left[ \begin{array}{c} a_1 \\ a_2 \\ 0 \end{array} \right] \  : \  a_1, a_2 \in \mathbb R\right\} .[/mathjax]</td>
        <td class="eqnnum" style="width:20%;vertical-align:middle;text-align:left;border-style:hidden"/>
      </tr>
    </table>
    <p>
 (In other words, you can think of [mathjaxinline]W[/mathjaxinline] as the [mathjaxinline]xy[/mathjaxinline] plane inside [mathjaxinline]\mathbb R^3[/mathjaxinline]. </p>
    <p>
Let [mathjaxinline]v = \left[ \begin{array}{c} 3 \\ 4 \\ 7 \end{array} \right].[/mathjaxinline] What is the vector [mathjaxinline]w\in W[/mathjaxinline] that is closest to [mathjaxinline]v[/mathjaxinline] (i.e., the vector such that the distance from [mathjaxinline]w[/mathjaxinline] to [mathjaxinline]v[/mathjaxinline] is smallest)? </p>
    <center>
      <img src="/static/images/c4s2projection.svg" width="300 px" style="margin: 10px 25px 25px 25px; border:0px"/>
    </center>
    <p>
      <div class="hideshowbox">
        <h4 onclick="hideshow(this);" style="margin: 0px">How to Enter a Column Vector in R^n<span class="icon-caret-down toggleimage"/></h4>
        <div class="hideshowcontent">
          <p>
To enter a vector in [mathjaxinline]\mathbb R^ n[/mathjaxinline] such as [mathjaxinline]\left[\begin{array}{c} 1 \\ 2 \\ 3 \end{array} \right][/mathjaxinline], you can either: </p>
          <ul class="itemize">
            <li>
              <p>
enter it as you would a [mathjaxinline]3\times 1[/mathjaxinline] matrix: [[1],[2],[3]] </p>
            </li>
            <li>
              <p>
or, you can enter it as &lt;1,2,3&gt; </p>
            </li>
          </ul>
          <p>
Use decimals only. </p>
        </div>
        <p class="hideshowbottom" onclick="hideshow(this);" style="margin: 0px">
          <a href="javascript: {return false;}">Show</a>
        </p>
      </div>
    </p>
    <script linenum="2" filename="./vectorentry2.tex" type="text/python" system_path="python_lib">
def VectorEntry(expect, ans):
	import ast
	import numpy as np 
  	atol = 0.01
	list_expect = ast.literal_eval(expect)
	vec_expect = np.matrix(list_expect)
  	ret = {"ok":False}
	try:
  		# input format [[1],[2],[3]]
		list_ans = ast.literal_eval(ans)
		vec_ans = np.matrix(list_ans)
  		if vec_ans.shape != vec_expect.shape:
  			ret['msg'] = 'Wrong shape of vector!'
  		elif np.allclose(vec_ans, vec_expect,atol,1e-08):
  			ret['ok'] = True
  		else:
  		# More error message. Will improve this part
  			ret['msg'] = 'something is wrong'   			
	#except SyntaxError:
		#ret['msg'] = 'Wrong input format'
	except SyntaxError:
  		# input format &amp;lt;1,2,3&amp;gt;
		list_ans = ans.replace('&amp;lt;', '[').replace('&amp;gt;', ']')
		list_ans = ast.literal_eval(list_ans)
		vec_ans = np.transpose(np.matrix(list_ans))
  		if vec_ans.shape != vec_expect.shape:
  			ret['msg'] = 'Wrong shape of vector!'
  		elif np.allclose(vec_ans, vec_expect,0.01,1e-08):
  			ret['ok'] = True
  		else:
    		# More error message. Will improve this part
  			ret['msg'] = 'something is wrong' 
  	except:
  		ret['msg'] = 'Wrong input format'
  	return ret 
</script>
    <p>
      <customresponse cfn="VectorEntry" expect="[[3],[4],[0]]">
        <textline correct_answer="[[3],[4],[0]]"/>
      </customresponse>
    </p>
    <p>
      <solution>
        <font color="blue">Answer: </font>
        <font color="blue">
          <p>
 If [mathjaxinline]w = \left[ \begin{array}{c} a_1 \\ a_2 \\ 0 \end{array} \right][/mathjaxinline], then the distance from [mathjaxinline]v[/mathjaxinline] to [mathjaxinline]w[/mathjaxinline] is </p>
          <table id="a0000000194" class="equation" width="100%" cellspacing="0" cellpadding="7" style="table-layout:auto">
            <tr>
              <td class="equation" style="width:80%;vertical-align:middle;text-align:center;border-style:hidden">[mathjax]\| w-v\|  = \sqrt {(a_1-3)^2 + (a_2-4)^2 + 7^2}.[/mathjax]</td>
              <td class="eqnnum" style="width:20%;vertical-align:middle;text-align:left;border-style:hidden"/>
            </tr>
          </table>
          <p>
 At minimum, this distance is [mathjaxinline]\sqrt {49} = 7[/mathjaxinline], and it equals 7 exactly when [mathjaxinline]a_1 = 3[/mathjaxinline] and [mathjaxinline]a_2=4[/mathjaxinline]. In other words, the distance is minimized when [mathjaxinline]w = \begin{array}{c} 3 \\ 4 \\ 0 \end{array} \right].[/mathjaxinline]  </p>
        </font>
      </solution>
    </p>
  </text>
  <SCRIPT src="/static/latex2edx.js" type="text/javascript"/>
  <LINK href="/static/latex2edx.css" rel="stylesheet" type="text/css"/>
</problem>
