<problem display_name="Another Closest Vector?" linenum="285" tex_filename="./block4/projections.tex" showanswer="finished" attempts="2" rerandomize="per_student">
  <text>
    <p>
Let [mathjaxinline]v = \left[ \begin{array}{c} 3 \\ 4 \\ 7 \\ 1 \end{array} \right].[/mathjaxinline] What is the vector [mathjaxinline]w\in \mathbb R^4[/mathjaxinline] that is closest to [mathjaxinline]v[/mathjaxinline]? </p>
    <p>
(You should not need to do any projection calculation!) </p>
    <script linenum="2" filename="./vectorentry2.tex" type="text/python" system_path="python_lib">
def VectorEntry(expect, ans):
	import ast
	import numpy as np 
  	atol = 0.01
	list_expect = ast.literal_eval(expect)
	vec_expect = np.matrix(list_expect)
  	ret = {"ok":False}
	try:
  		# input format [[1],[2],[3]]
		list_ans = ast.literal_eval(ans)
		vec_ans = np.matrix(list_ans)
  		if vec_ans.shape != vec_expect.shape:
  			ret['msg'] = 'Wrong shape of vector!'
  		elif np.allclose(vec_ans, vec_expect,atol,1e-08):
  			ret['ok'] = True
  		else:
  		# More error message. Will improve this part
  			ret['msg'] = 'something is wrong'   			
	#except SyntaxError:
		#ret['msg'] = 'Wrong input format'
	except SyntaxError:
  		# input format &amp;lt;1,2,3&amp;gt;
		list_ans = ans.replace('&amp;lt;', '[').replace('&amp;gt;', ']')
		list_ans = ast.literal_eval(list_ans)
		vec_ans = np.transpose(np.matrix(list_ans))
  		if vec_ans.shape != vec_expect.shape:
  			ret['msg'] = 'Wrong shape of vector!'
  		elif np.allclose(vec_ans, vec_expect,0.01,1e-08):
  			ret['ok'] = True
  		else:
    		# More error message. Will improve this part
  			ret['msg'] = 'something is wrong' 
  	except:
  		ret['msg'] = 'Wrong input format'
  	return ret 
</script>
    <p>
      <customresponse cfn="VectorEntry" expect="[[3],[4],[7],[1]]">
        <textline correct_answer="[[3],[4],[7],[1]]"/>
      </customresponse>
    </p>
    <p>
      <solution>
        <font color="blue">Answer: </font>
        <font color="blue">
          <p>
 Certainly [mathjaxinline]v[/mathjaxinline] itself is the only vector that is distance zero from [mathjaxinline]v[/mathjaxinline]!  </p>
        </font>
      </solution>
    </p>
  </text>
</problem>
