<problem display_name="Gram-Schmidt Practice 3" linenum="399" tex_filename="./block4/projections.tex" showanswer="finished" attempts="3" rerandomize="per_student">
  <text>
    <p>
Finish the algorithm to calculate [mathjaxinline]w_3[/mathjaxinline]. </p>
    <script linenum="2" filename="./vectorentry2.tex" type="text/python" system_path="python_lib">
def VectorEntry(expect, ans):
	import ast
	import numpy as np 
  	atol = 0.01
	list_expect = ast.literal_eval(expect)
	vec_expect = np.matrix(list_expect)
  	ret = {"ok":False}
	try:
  		# input format [[1],[2],[3]]
		list_ans = ast.literal_eval(ans)
		vec_ans = np.matrix(list_ans)
  		if vec_ans.shape != vec_expect.shape:
  			ret['msg'] = 'Wrong shape of vector!'
  		elif np.allclose(vec_ans, vec_expect,atol,1e-08):
  			ret['ok'] = True
  		else:
  		# More error message. Will improve this part
  			ret['msg'] = 'something is wrong'   			
	#except SyntaxError:
		#ret['msg'] = 'Wrong input format'
	except SyntaxError:
  		# input format &amp;lt;1,2,3&amp;gt;
		list_ans = ans.replace('&amp;lt;', '[').replace('&amp;gt;', ']')
		list_ans = ast.literal_eval(list_ans)
		vec_ans = np.transpose(np.matrix(list_ans))
  		if vec_ans.shape != vec_expect.shape:
  			ret['msg'] = 'Wrong shape of vector!'
  		elif np.allclose(vec_ans, vec_expect,0.01,1e-08):
  			ret['ok'] = True
  		else:
    		# More error message. Will improve this part
  			ret['msg'] = 'something is wrong' 
  	except:
  		ret['msg'] = 'Wrong input format'
  	return ret 
</script>
    <p>
      <customresponse cfn="VectorEntry" expect="[[1],[-3],[3],[-1]]">
        <textline correct_answer="[[1],[-3],[3],[-1]]"/>
      </customresponse>
    </p>
    <p>
      <solution>
        <font color="blue">Answer: </font>
        <font color="blue">
          <p>
 Again, [mathjaxinline]w_3[/mathjaxinline] will be [mathjaxinline]u_3[/mathjaxinline] minus its projection onto the span of [mathjaxinline]w_1[/mathjaxinline] and [mathjaxinline]w_2[/mathjaxinline] (which we can call [mathjaxinline]W_2[/mathjaxinline]). The projection is </p>
          <table id="a0000000202" class="equation" width="100%" cellspacing="0" cellpadding="7" style="table-layout:auto">
            <tr>
              <td class="equation" style="width:80%;vertical-align:middle;text-align:center;border-style:hidden">[mathjax]\mathrm{Proj}_{W_2}(u_3) = \frac{\langle u_3, w_1 \rangle }{\langle w_1, w_1 \rangle } w_1 + \frac{\langle u_3, w_2 \rangle }{\langle w_2, w_2 \rangle } w_2 = \frac{8}{4}w_1 + \frac{-14}{14}w_2 =\left[ \begin{array}{c} -1 \\ -4 \\ -3 \\ 2 \end{array} \right],[/mathjax]</td>
              <td class="eqnnum" style="width:20%;vertical-align:middle;text-align:left;border-style:hidden"/>
            </tr>
          </table>
          <p>
 so [mathjaxinline]w_3 = u_3 -\mathrm{Proj}_{W_2}(u_3) = \left[ \begin{array}{c} 1 \\ -3 \\ 3 \\ -1 \end{array} \right].[/mathjaxinline]  </p>
        </font>
      </solution>
    </p>
  </text>
</problem>
