<problem display_name="Invert a Matrix" linenum="196" tex_filename="./hw/final.tex" showanswer="finished" attempts="3" rerandomize="per_student">
  <text>
    <p>
Let [mathjaxinline]A[/mathjaxinline] be the matrix </p>
    <table id="a0000000221" class="equation" width="100%" cellspacing="0" cellpadding="7" style="table-layout:auto">
      <tr>
        <td class="equation" style="width:80%;vertical-align:middle;text-align:center;border-style:hidden">[mathjax]\left[ \begin{array}{ccc} 1&amp;  1 &amp;  2 \\ 2 &amp;  1 &amp;  -3 \\ 1 &amp;  -1 &amp;  -11 \end{array} \right][/mathjax]</td>
        <td class="eqnnum" style="width:20%;vertical-align:middle;text-align:left;border-style:hidden"/>
      </tr>
    </table>
    <p>
Enter its inverse, if it has one. If it does not have an inverse, enter [[0]]. </p>
    <script linenum="1" filename="./matrixentry3.tex" type="text/python" system_path="python_lib">
def MatrixEntry(expect, ans):
  	import ast
	import numpy as np 
	ret= {'ok':False}
  	atol = 0.01
  	try:
		list_ans = ast.literal_eval(ans)
		list_expect = ast.literal_eval(expect)
  		matrix_ans = np.matrix(list_ans)
  		matrix_expect = np.matrix(list_expect) 
  		if matrix_ans.shape != matrix_expect.shape:
  			ret['msg'] = 'Something is wrong'
  		elif np.allclose(matrix_ans, matrix_expect,0.01,1e-08):
  			ret['ok'] = True
  		else:
  			ret['msg'] = 'Something is wrong'
	except SyntaxError:
		ret['msg'] = 'Wrong input format'
  	return ret
</script>
    <p>
      <customresponse cfn="MatrixEntry" expect="[[14,-9,5],[-19,13,-7],[3,-2,1]]">
        <textline correct_answer="[[14,-9,5],[-19,13,-7],[3,-2,1]]"/>
      </customresponse>
    </p>
    <p>
      <solution>
        <font color="blue">Answer: </font>
        <font color="blue">
          <p>
We can check that [mathjaxinline]A[/mathjaxinline] row reduces to the identity. When we row reduce the extended </p>
          <table id="a0000000222" class="equation" width="100%" cellspacing="0" cellpadding="7" style="table-layout:auto">
            <tr>
              <td class="equation" style="width:80%;vertical-align:middle;text-align:center;border-style:hidden">[mathjax]\left[ \begin{array}{ccc:ccc} 1&amp;  1 &amp;  2 &amp;  1 &amp;  0 &amp;  0 \\ 2 &amp;  1 &amp;  -3 &amp;  0 &amp;  1 &amp;  0 \\ 1 &amp;  -1 &amp;  -11 &amp;  0 &amp;  0 &amp;  1 \end{array} \right],[/mathjax]</td>
              <td class="eqnnum" style="width:20%;vertical-align:middle;text-align:left;border-style:hidden"/>
            </tr>
          </table>
          <p>
 we obtain </p>
          <table id="a0000000223" class="equation" width="100%" cellspacing="0" cellpadding="7" style="table-layout:auto">
            <tr>
              <td class="equation" style="width:80%;vertical-align:middle;text-align:center;border-style:hidden">[mathjax]\left[ \begin{array}{ccc:ccc} 1&amp;  0 &amp;  0 &amp;  14 &amp;  -9 &amp;  5 \\ 0 &amp;  1 &amp;  0 &amp;  -19 &amp;  13 &amp;  -7 \\ 0 &amp;  0 &amp;  1 &amp;  3 &amp;  -2 &amp;  1 \end{array} \right],[/mathjax]</td>
              <td class="eqnnum" style="width:20%;vertical-align:middle;text-align:left;border-style:hidden"/>
            </tr>
          </table>
          <p>
 so the inverse is </p>
          <table id="a0000000224" class="equation" width="100%" cellspacing="0" cellpadding="7" style="table-layout:auto">
            <tr>
              <td class="equation" style="width:80%;vertical-align:middle;text-align:center;border-style:hidden">[mathjax]\left[ \begin{array}{ccc} 14 &amp;  -9 &amp;  5 \\ -19 &amp;  13 &amp;  -7 \\ 3 &amp;  -2 &amp;  1 \end{array} \right].[/mathjax]</td>
              <td class="eqnnum" style="width:20%;vertical-align:middle;text-align:left;border-style:hidden"/>
            </tr>
          </table>
        </font>
      </solution>
    </p>
  </text>
</problem>
