<problem display_name="Linear Transformations" linenum="145" tex_filename="./hw/final.tex" showanswer="finished" attempts="3" rerandomize="per_student">
  <text>
    <p>
Let [mathjaxinline]S: \mathbb R^3 \rightarrow \mathbb R^2[/mathjaxinline] be a linear transformation with </p>
    <table id="a0000000220" class="equation" width="100%" cellspacing="0" cellpadding="7" style="table-layout:auto">
      <tr>
        <td class="equation" style="width:80%;vertical-align:middle;text-align:center;border-style:hidden">[mathjax]S(e_1) = \left[ \begin{array}{c} -1 \\ 4 \end{array} \right], S(e_2) = \left[ \begin{array}{c} 2 \\ 1 \end{array} \right], S(e_3) = \left[ \begin{array}{c} 3 \\ 3 \end{array} \right].[/mathjax]</td>
        <td class="eqnnum" style="width:20%;vertical-align:middle;text-align:left;border-style:hidden"/>
      </tr>
    </table>
    <p>
Let [mathjaxinline]T: \mathbb R^2 \rightarrow \mathbb R^4[/mathjaxinline] be a linear transformation with [mathjaxinline]T(e_1) = \left[ \begin{array}{c} 3 \\ -1 \\ 0 \\ 1 \end{array} \right][/mathjaxinline] and [mathjaxinline]T(e_2) = \left[ \begin{array}{c} 1 \\ 0 \\ 1 \\ -2 \end{array} \right].[/mathjaxinline] </p>
    <p>
What is the standard matrix for [mathjaxinline]T\circ S[/mathjaxinline]? </p>
    <p>
      <div class="hideshowbox">
        <h4 onclick="hideshow(this);" style="margin: 0px">How to Enter a Matrix<span class="icon-caret-down toggleimage"/></h4>
        <div class="hideshowcontent">
          <p>
To enter the matrix [mathjaxinline]\left[ \begin{array}{ccc} 0&amp; 1&amp; 2 \\ 1&amp; 2&amp; 3 \end{array} \right][/mathjaxinline], type [[0,1,2],[1,2,3]]. (Type the same thing if you want to enter the augmented matrix [mathjaxinline]\left[ \begin{array}{cc:c} 0&amp; 1&amp; 2 \\ 1&amp; 2&amp; 3 \end{array} \right][/mathjaxinline].) </p>
          <p>
Use decimals only. </p>
        </div>
        <p class="hideshowbottom" onclick="hideshow(this);" style="margin: 0px">
          <a href="javascript: {return false;}">Show</a>
        </p>
      </div>
    </p>
    <script linenum="1" filename="./matrixentry2.tex" type="text/python" system_path="python_lib">
def MatrixEntry(expect, ans):
  	import ast
	import numpy as np 
	ret= {'ok':False}
  	atol = 0.01
  	try:
		list_ans = ast.literal_eval(ans)
		list_expect = ast.literal_eval(expect)
  		matrix_ans = np.matrix(list_ans)
  		matrix_expect = np.matrix(list_expect) 
  		if matrix_ans.shape != matrix_expect.shape:
  			ret['msg'] = 'Wrong shape of matrix'
  		elif np.allclose(matrix_ans, matrix_expect,0.01,1e-08):
  			ret['ok'] = True
  		else:
  			ret['msg'] = 'Something is wrong'
	except SyntaxError:
		ret['msg'] = 'Wrong input format'
  	return ret
</script>
    <p>
      <customresponse cfn="MatrixEntry" expect="[[1,7,12],[1,-2,-3],[4,1,3],[-9,0,-3]]">
        <textline correct_answer="[[1,7,12],[1,-2,-3],[4,1,3],[-9,0,-3]]"/>
      </customresponse>
    </p>
    <p>
      <solution>
        <font color="blue">Answer: </font>
        <font color="blue"/>
      </solution>
    </p>
  </text>
  <SCRIPT src="/static/latex2edx.js" type="text/javascript"/>
  <LINK href="/static/latex2edx.css" rel="stylesheet" type="text/css"/>
</problem>
