<problem display_name="Matrix-Vector Multiplication" linenum="279" tex_filename="./hw/hw1.tex" showanswer="finished" attempts="3" rerandomize="per_student">
  <text>
    <p>
Let [mathjaxinline]u,v,w[/mathjaxinline] be three vectors in [mathjaxinline]\mathbb R^4[/mathjaxinline] with the property that [mathjaxinline]4u - 3v + 2w = \mathbf{0}[/mathjaxinline]. Let [mathjaxinline]A[/mathjaxinline] be the [mathjaxinline]4\times 2[/mathjaxinline] matrix whose columns are [mathjaxinline]u[/mathjaxinline] and [mathjaxinline]v[/mathjaxinline] (in that order). </p>
    <p>
Find a solution to the equation [mathjaxinline]Ax = w[/mathjaxinline]. </p>
    <p>
      <div class="hideshowbox">
        <h4 onclick="hideshow(this);" style="margin: 0px">How to Enter a Column Vector in R^n<span class="icon-caret-down toggleimage"/></h4>
        <div class="hideshowcontent">
          <p>
To enter a vector in [mathjaxinline]\mathbb R^ n[/mathjaxinline] such as [mathjaxinline]\left[\begin{array}{c} 1 \\ 2 \\ 3 \end{array} \right][/mathjaxinline], you can either: </p>
          <ul class="itemize">
            <li>
              <p>
enter it as you would a [mathjaxinline]3\times 1[/mathjaxinline] matrix: [[1],[2],[3]] </p>
            </li>
            <li>
              <p>
or, you can enter it as &lt;1,2,3&gt; </p>
            </li>
          </ul>
          <p>
Use decimals only. </p>
        </div>
        <p class="hideshowbottom" onclick="hideshow(this);" style="margin: 0px">
          <a href="javascript: {return false;}">Show</a>
        </p>
      </div>
    </p>
    <script linenum="2" filename="./vectorentry2.tex" type="text/python" system_path="python_lib">
def VectorEntry(expect, ans):
	import ast
	import numpy as np 
  	atol = 0.01
	list_expect = ast.literal_eval(expect)
	vec_expect = np.matrix(list_expect)
  	ret = {"ok":False}
	try:
  		# input format [[1],[2],[3]]
		list_ans = ast.literal_eval(ans)
		vec_ans = np.matrix(list_ans)
  		if vec_ans.shape != vec_expect.shape:
  			ret['msg'] = 'Wrong shape of vector!'
  		elif np.allclose(vec_ans, vec_expect,atol,1e-08):
  			ret['ok'] = True
  		else:
  		# More error message. Will improve this part
  			ret['msg'] = 'something is wrong'   			
	#except SyntaxError:
		#ret['msg'] = 'Wrong input format'
	except SyntaxError:
  		# input format &amp;lt;1,2,3&amp;gt;
		list_ans = ans.replace('&amp;lt;', '[').replace('&amp;gt;', ']')
		list_ans = ast.literal_eval(list_ans)
		vec_ans = np.transpose(np.matrix(list_ans))
  		if vec_ans.shape != vec_expect.shape:
  			ret['msg'] = 'Wrong shape of vector!'
  		elif np.allclose(vec_ans, vec_expect,0.01,1e-08):
  			ret['ok'] = True
  		else:
    		# More error message. Will improve this part
  			ret['msg'] = 'something is wrong' 
  	except:
  		ret['msg'] = 'Wrong input format'
  	return ret 
</script>
    <p>
      <customresponse cfn="VectorEntry" expect="[[-2],[1.5]]">
        <textline correct_answer="[[-2],[1.5]]"/>
      </customresponse>
    </p>
    <p>
      <solution>
        <font color="blue">Answer: </font>
        <font color="blue">
          <p>
 We can solve the given equation for [mathjaxinline]w[/mathjaxinline]: </p>
          <table id="a0000000101" class="equation" width="100%" cellspacing="0" cellpadding="7" style="table-layout:auto">
            <tr>
              <td class="equation" style="width:80%;vertical-align:middle;text-align:center;border-style:hidden">[mathjax]w = -2u + 1.5v.[/mathjax]</td>
              <td class="eqnnum" style="width:20%;vertical-align:middle;text-align:left;border-style:hidden"/>
            </tr>
          </table>
          <p>
 This is a linear combination of the columns of [mathjaxinline]A[/mathjaxinline]; it is [mathjaxinline]-2[/mathjaxinline] times the first column of [mathjaxinline]A[/mathjaxinline] plus [mathjaxinline]1.5[/mathjaxinline] times the second column of [mathjaxinline]A[/mathjaxinline]. By definition, this is the same as the matrix-vector product [mathjaxinline]A\left[ \begin{array}{c} -2 \\ 1.5 \end{array}\right].[/mathjaxinline] </p>
          <p>
Therefore [mathjaxinline]x = \left[ \begin{array}{c} -2 \\ 1.5 \end{array}\right][/mathjaxinline] is a solution to [mathjaxinline]Ax = w.[/mathjaxinline]  </p>
        </font>
      </solution>
    </p>
  </text>
  <SCRIPT src="/static/latex2edx.js" type="text/javascript"/>
  <LINK href="/static/latex2edx.css" rel="stylesheet" type="text/css"/>
</problem>
