<problem display_name="Invert a Matrix" linenum="102" tex_filename="./hw/hw2.tex" showanswer="finished" attempts="3" rerandomize="per_student">
  <text>
    <p>
For each of the following matrices, enter its inverse, if it has one. If it does not have an inverse, enter [[0]]. </p>
    <script linenum="1" filename="./matrixentry3.tex" type="text/python" system_path="python_lib">
def MatrixEntry(expect, ans):
  	import ast
	import numpy as np 
	ret= {'ok':False}
  	atol = 0.01
  	try:
		list_ans = ast.literal_eval(ans)
		list_expect = ast.literal_eval(expect)
  		matrix_ans = np.matrix(list_ans)
  		matrix_expect = np.matrix(list_expect) 
  		if matrix_ans.shape != matrix_expect.shape:
  			ret['msg'] = 'Something is wrong'
  		elif np.allclose(matrix_ans, matrix_expect,0.01,1e-08):
  			ret['ok'] = True
  		else:
  			ret['msg'] = 'Something is wrong'
	except SyntaxError:
		ret['msg'] = 'Wrong input format'
  	return ret
</script>
    <p>
[mathjaxinline]\left[ \begin{array}{ccc} 1&amp;  2 &amp;  -1 \\ 1 &amp;  0 &amp;  2 \end{array} \right][/mathjaxinline] </p>
    <p>
      <customresponse cfn="MatrixEntry" expect="[[0]]">
        <textline correct_answer="[[0]]"/>
      </customresponse>
    </p>
    <p>
[mathjaxinline]\left[ \begin{array}{ccc} 1&amp;  3 &amp;  -1 \\ 2 &amp;  2 &amp;  1 \\ 0 &amp;  4 &amp;  -3 \end{array} \right][/mathjaxinline] </p>
    <p>
      <customresponse cfn="MatrixEntry" expect="[[0]]">
        <textline correct_answer="[[0]]"/>
      </customresponse>
    </p>
    <p>
[mathjaxinline]\left[ \begin{array}{ccc} 1&amp;  1 &amp;  0 \\ 1 &amp;  2 &amp;  -3 \\ 1 &amp;  1 &amp;  1 \end{array} \right][/mathjaxinline] </p>
    <p>
      <customresponse cfn="MatrixEntry" expect="[[5,-1,-3],[-4,1,3],[-1,0,1]]">
        <textline correct_answer="[[5,-1,-3],[-4,1,3],[-1,0,1]]"/>
      </customresponse>
    </p>
    <p>
      <solution>
        <font color="blue">Answer: </font>
        <font color="blue">
          <p>
 The first matrix is not square, so it has no inverse. </p>
          <p>
The second matrix is square. However, when we row reduce it, we do not get the identity matrix. Therefore it too is non-invertible. </p>
          <p>
The last matrix does row reduce to the identity. When we row reduce the extended </p>
          <table id="a0000000136" class="equation" width="100%" cellspacing="0" cellpadding="7" style="table-layout:auto">
            <tr>
              <td class="equation" style="width:80%;vertical-align:middle;text-align:center;border-style:hidden">[mathjax]\left[ \begin{array}{ccc:ccc} 1&amp;  1 &amp;  0 &amp;  1 &amp;  0 &amp;  0 \\ 1 &amp;  2 &amp;  -3 &amp;  0 &amp;  1 &amp;  0 \\ 1 &amp;  1 &amp;  1 &amp;  0 &amp;  0 &amp;  1 \end{array} \right],[/mathjax]</td>
              <td class="eqnnum" style="width:20%;vertical-align:middle;text-align:left;border-style:hidden"/>
            </tr>
          </table>
          <p>
 we obtain </p>
          <table id="a0000000137" class="equation" width="100%" cellspacing="0" cellpadding="7" style="table-layout:auto">
            <tr>
              <td class="equation" style="width:80%;vertical-align:middle;text-align:center;border-style:hidden">[mathjax]\left[ \begin{array}{ccc:ccc} 1&amp;  0 &amp;  0 &amp;  5 &amp;  -1 &amp;  -3 \\ 0 &amp;  1 &amp;  0 &amp;  -4 &amp;  1 &amp;  3 \\ 0 &amp;  0 &amp;  1 &amp;  -1 &amp;  0 &amp;  1 \end{array} \right],[/mathjax]</td>
              <td class="eqnnum" style="width:20%;vertical-align:middle;text-align:left;border-style:hidden"/>
            </tr>
          </table>
          <p>
 so the inverse is </p>
          <table id="a0000000138" class="equation" width="100%" cellspacing="0" cellpadding="7" style="table-layout:auto">
            <tr>
              <td class="equation" style="width:80%;vertical-align:middle;text-align:center;border-style:hidden">[mathjax]\left[ \begin{array}{ccc} 5 &amp;  -1 &amp;  -3 \\ -4 &amp;  1 &amp;  3 \\ -1 &amp;  0 &amp;  1 \end{array} \right].[/mathjax]</td>
              <td class="eqnnum" style="width:20%;vertical-align:middle;text-align:left;border-style:hidden"/>
            </tr>
          </table>
        </font>
      </solution>
    </p>
  </text>
</problem>
