<problem display_name="Linear Transformation" linenum="253" tex_filename="./hw/hw2.tex" showanswer="finished" attempts="4" rerandomize="per_student">
  <text>
    <p>
Suppose that [mathjaxinline]T:\mathbb R^3\rightarrow \mathbb R^3[/mathjaxinline] is a linear transformation with the property that </p>
    <table id="a0000000140" class="equation" width="100%" cellspacing="0" cellpadding="7" style="table-layout:auto">
      <tr>
        <td class="equation" style="width:80%;vertical-align:middle;text-align:center;border-style:hidden">[mathjax]\begin{array}{rcl} T\left(\left[ \begin{array}{c} 1 \\ 7 \\ 0\end{array} \right]\right) &amp; =&amp;  e_1\\ T\left(\left[ \begin{array}{c} 3 \\ 2 \\ 1\end{array} \right]\right) &amp; =&amp;  e_2\\ T\left(\left[ \begin{array}{c} 0 \\ 5 \\ 3\end{array} \right]\right) &amp; =&amp;  e_3. \end{array}[/mathjax]</td>
        <td class="eqnnum" style="width:20%;vertical-align:middle;text-align:left;border-style:hidden"/>
      </tr>
    </table>
    <p>
Find a vector [mathjaxinline]w[/mathjaxinline] such that [mathjaxinline]T(w) = \left[ \begin{array}{c} 1 \\ 3 \\ -2\end{array} \right][/mathjaxinline]. </p>
    <script linenum="2" filename="./vectorentry2.tex" type="text/python" system_path="python_lib">
def VectorEntry(expect, ans):
	import ast
	import numpy as np 
  	atol = 0.01
	list_expect = ast.literal_eval(expect)
	vec_expect = np.matrix(list_expect)
  	ret = {"ok":False}
	try:
  		# input format [[1],[2],[3]]
		list_ans = ast.literal_eval(ans)
		vec_ans = np.matrix(list_ans)
  		if vec_ans.shape != vec_expect.shape:
  			ret['msg'] = 'Wrong shape of vector!'
  		elif np.allclose(vec_ans, vec_expect,atol,1e-08):
  			ret['ok'] = True
  		else:
  		# More error message. Will improve this part
  			ret['msg'] = 'something is wrong'   			
	#except SyntaxError:
		#ret['msg'] = 'Wrong input format'
	except SyntaxError:
  		# input format &amp;lt;1,2,3&amp;gt;
		list_ans = ans.replace('&amp;lt;', '[').replace('&amp;gt;', ']')
		list_ans = ast.literal_eval(list_ans)
		vec_ans = np.transpose(np.matrix(list_ans))
  		if vec_ans.shape != vec_expect.shape:
  			ret['msg'] = 'Wrong shape of vector!'
  		elif np.allclose(vec_ans, vec_expect,0.01,1e-08):
  			ret['ok'] = True
  		else:
    		# More error message. Will improve this part
  			ret['msg'] = 'something is wrong' 
  	except:
  		ret['msg'] = 'Wrong input format'
  	return ret 
</script>
    <p>
      <customresponse cfn="VectorEntry" expect="[[10],[3],[-3]]">
        <textline correct_answer="[[10],[3],[-3]]"/>
      </customresponse>
    </p>
    <p>
Find a vector [mathjaxinline]w[/mathjaxinline] such that [mathjaxinline]T(w) = \left[ \begin{array}{c} 1000 \\ 100 \\ 10 \end{array} \right][/mathjaxinline]. </p>
    <p>
      <customresponse cfn="VectorEntry" expect="[[1300],[7250],[130]]">
        <textline correct_answer="[[1300],[7250],[130]]"/>
      </customresponse>
    </p>
    <p>
      <solution>
        <font color="blue">Answer: </font>
        <font color="blue">
          <p>
 Call the three given vectors [mathjaxinline]v_1, v_2, v_3,[/mathjaxinline] such that [mathjaxinline]T(v_ i) = e_ i[/mathjaxinline]. Note that </p>
          <table id="a0000000141" class="equation" width="100%" cellspacing="0" cellpadding="7" style="table-layout:auto">
            <tr>
              <td class="equation" style="width:80%;vertical-align:middle;text-align:center;border-style:hidden">[mathjax]\left[ \begin{array}{c} 1 \\ 3 \\ -2\end{array} \right] = 1e_1 + 3e_2 - 2e_3 = T(v_1 + 3v_2 - 2v_3).[/mathjax]</td>
              <td class="eqnnum" style="width:20%;vertical-align:middle;text-align:left;border-style:hidden"/>
            </tr>
          </table>
          <p>
 Therefore </p>
          <table id="a0000000142" class="equation" width="100%" cellspacing="0" cellpadding="7" style="table-layout:auto">
            <tr>
              <td class="equation" style="width:80%;vertical-align:middle;text-align:center;border-style:hidden">[mathjax]w = v_1 + 3v_2 - 2v_3 = \left[ \begin{array}{c} 10 \\ 3 \\ -3\end{array} \right][/mathjax]</td>
              <td class="eqnnum" style="width:20%;vertical-align:middle;text-align:left;border-style:hidden"/>
            </tr>
          </table>
          <p>
 is a solution to the first equation. </p>
          <p>
Similarly, </p>
          <table id="a0000000143" class="equation" width="100%" cellspacing="0" cellpadding="7" style="table-layout:auto">
            <tr>
              <td class="equation" style="width:80%;vertical-align:middle;text-align:center;border-style:hidden">[mathjax]w = 1000v_1 + 100v_2 + 10v_3 = \left[ \begin{array}{c} 1300 \\ 7250 \\ 130\end{array} \right][/mathjax]</td>
              <td class="eqnnum" style="width:20%;vertical-align:middle;text-align:left;border-style:hidden"/>
            </tr>
          </table>
          <p>
 is a solution to the second equation.  </p>
        </font>
      </solution>
    </p>
  </text>
</problem>
