<problem display_name="Solve" linenum="154" tex_filename="./hw/hw2.tex" showanswer="finished" attempts="4" rerandomize="per_student">
  <text>
    <p>
Suppose that [mathjaxinline]T:\mathbb R^4\rightarrow \mathbb R^4[/mathjaxinline] is a linear transformation with standard matrix [mathjaxinline]A[/mathjaxinline]. Suppose further that we know that [mathjaxinline]A[/mathjaxinline] is invertible, and that that </p>
    <table id="a0000000139" class="equation" width="100%" cellspacing="0" cellpadding="7" style="table-layout:auto">
      <tr>
        <td class="equation" style="width:80%;vertical-align:middle;text-align:center;border-style:hidden">[mathjax]A^{-1}= \left[ \begin{array}{cccc} 1&amp;  1 &amp;  0 &amp;  a_1 \\ 1 &amp;  2 &amp;  -3 &amp;  a_2 \\ 1 &amp;  1 &amp;  1 &amp;  a_3 \\ 7 &amp;  0 &amp;  1 &amp;  a_4 \end{array} \right],[/mathjax]</td>
        <td class="eqnnum" style="width:20%;vertical-align:middle;text-align:left;border-style:hidden"/>
      </tr>
    </table>
    <p>
 for some scalars [mathjaxinline]a_1, a_2, a_3, a_4[/mathjaxinline]. </p>
    <p>
Solve the equation [mathjaxinline]T(v) = \left[ \begin{array}{c} 3 \\ 2 \\ 1 \\ 0 \end{array} \right][/mathjaxinline]. </p>
    <p>
      <div class="hideshowbox">
        <h4 onclick="hideshow(this);" style="margin: 0px">How to Enter a Column Vector in R^n<span class="icon-caret-down toggleimage"/></h4>
        <div class="hideshowcontent">
          <p>
To enter a vector in [mathjaxinline]\mathbb R^ n[/mathjaxinline] such as [mathjaxinline]\left[\begin{array}{c} 1 \\ 2 \\ 3 \end{array} \right][/mathjaxinline], you can either: </p>
          <ul class="itemize">
            <li>
              <p>
enter it as you would a [mathjaxinline]3\times 1[/mathjaxinline] matrix: [[1],[2],[3]] </p>
            </li>
            <li>
              <p>
or, you can enter it as &lt;1,2,3&gt; </p>
            </li>
          </ul>
          <p>
Use decimals only. </p>
        </div>
        <p class="hideshowbottom" onclick="hideshow(this);" style="margin: 0px">
          <a href="javascript: {return false;}">Show</a>
        </p>
      </div>
    </p>
    <script linenum="2" filename="./vectorentry2.tex" type="text/python" system_path="python_lib">
def VectorEntry(expect, ans):
	import ast
	import numpy as np 
  	atol = 0.01
	list_expect = ast.literal_eval(expect)
	vec_expect = np.matrix(list_expect)
  	ret = {"ok":False}
	try:
  		# input format [[1],[2],[3]]
		list_ans = ast.literal_eval(ans)
		vec_ans = np.matrix(list_ans)
  		if vec_ans.shape != vec_expect.shape:
  			ret['msg'] = 'Wrong shape of vector!'
  		elif np.allclose(vec_ans, vec_expect,atol,1e-08):
  			ret['ok'] = True
  		else:
  		# More error message. Will improve this part
  			ret['msg'] = 'something is wrong'   			
	#except SyntaxError:
		#ret['msg'] = 'Wrong input format'
	except SyntaxError:
  		# input format &amp;lt;1,2,3&amp;gt;
		list_ans = ans.replace('&amp;lt;', '[').replace('&amp;gt;', ']')
		list_ans = ast.literal_eval(list_ans)
		vec_ans = np.transpose(np.matrix(list_ans))
  		if vec_ans.shape != vec_expect.shape:
  			ret['msg'] = 'Wrong shape of vector!'
  		elif np.allclose(vec_ans, vec_expect,0.01,1e-08):
  			ret['ok'] = True
  		else:
    		# More error message. Will improve this part
  			ret['msg'] = 'something is wrong' 
  	except:
  		ret['msg'] = 'Wrong input format'
  	return ret 
</script>
    <p>
      <customresponse cfn="VectorEntry" expect="[[5],[4],[6],[22]]">
        <textline correct_answer="[[5],[4],[6],[22]]"/>
      </customresponse>
    </p>
    <p>
      <solution>
        <font color="blue">Answer: </font>
        <font color="blue">
          <p>
 Solving [mathjaxinline]T(v) = w[/mathjaxinline] is equivalent to solving [mathjaxinline]Av = w[/mathjaxinline]; since [mathjaxinline]A[/mathjaxinline] is invertible we know the solution must be [mathjaxinline]v = A^{-1}w[/mathjaxinline]. Thus, we multiply [mathjaxinline]A^{-1}[/mathjaxinline] by [mathjaxinline]\left[ \begin{array}{c} 3 \\ 2 \\ 1 \\ 0 \end{array} \right],[/mathjaxinline] and obtain [mathjaxinline]\left[ \begin{array}{c} 5 \\ 4 \\ 6 \\ 22 \end{array} \right].[/mathjaxinline] (It does not matter that we don't know the [mathjaxinline]a_ i[/mathjaxinline] since they get multiplied by zero.)  </p>
        </font>
      </solution>
    </p>
  </text>
  <SCRIPT src="/static/latex2edx.js" type="text/javascript"/>
  <LINK href="/static/latex2edx.css" rel="stylesheet" type="text/css"/>
</problem>
