<problem display_name="Kernel Basis" linenum="311" tex_filename="./hw/hw3.tex" showanswer="finished" attempts="5" rerandomize="per_student">
  <text>
    <p>
Suppose [mathjaxinline]T: \mathbb R^4 \rightarrow \mathbb R^4[/mathjaxinline] is a linear transformation with standard matrix </p>
    <table id="a0000000174" class="equation" width="100%" cellspacing="0" cellpadding="7" style="table-layout:auto">
      <tr>
        <td class="equation" style="width:80%;vertical-align:middle;text-align:center;border-style:hidden">[mathjax]A = \left[ \begin{array}{cccc} 3 &amp;  1 &amp;  1 &amp;  -2 \\ -1 &amp;  -2 &amp;  3 &amp;  -1 \\ 1 &amp;  -3 &amp;  7 &amp;  -4 \\ -3 &amp;  4 &amp;  -11 &amp;  7 \end{array} \right][/mathjax]</td>
        <td class="eqnnum" style="width:20%;vertical-align:middle;text-align:left;border-style:hidden"/>
      </tr>
    </table>
    <p>
Find a basis for the kernel of [mathjaxinline]T[/mathjaxinline]. Enter the list of vectors forming the basis below, separated by semicolons. For instance, to enter the list [mathjaxinline]\left\{ \left[\begin{array}{c} 1 \\ 0 \\ 1 \end{array} \right]; \left[\begin{array}{c} 1 \\ 2 \\ 3 \end{array} \right] \right\}[/mathjaxinline], type &lt;1,0,1&gt;;&lt;1,2,3&gt;. </p>
    <script linenum="2" filename="./solutionsetcheck3.tex" type="text/python" system_path="python_lib">

def VectorEntry(ans):
    import numpy as np
    import ast
    test={'ok':False}
    try:
        ans=ans.split(";")
        hold=[]
        for a in ans:
            a=a.split(",")
            sub=[]
            for i in range(len(a)):
                if i==0:
                    sub.append(float(a[i][1:]))
                elif i==len(a)-1:
                    sub.append(float(a[i][:-1]))
                else:
                    sub.append(float(a[i]))
            assert len(sub)==4
            hold.append(np.array(sub))
        if len(hold)==2:
            test['ok']=True
        else:
            test['msg']='You gave the wrong number of vectors'        
    except AssertionError:
        test['msg']='One of your vectors is not a 4-tuple'
    except:
        test['msg']='Wrong input format'
    return test  
def span2(a,b):
    import numpy as np
    import ast
    non=np.nonzero(a)
    first=non[0][0]
    den=a[first]
    num=b[first]
    quot=num//den
    if np.array_equal(np.multiply(a,quot),b):
        return(-1)
    else:
        return(1)
        
def spangrader(expect,ans):
    import numpy as np
    import ast
    mata=np.array([[3,1,1,-2],[-1,-2,3,-1],[1,-3,7,-4],[-3,4,-11,7]])
    res={'ok':False}
    if VectorEntry(ans)['ok']!=True:
        return(VectorEntry(ans))
    else:
        ans=ans.split(";")
        hold=[]
        for a in ans:
            a=a.split(",")
            sub=[]
            for i in range(len(a)):
                if i==0:
                    sub.append(float(a[i][1:]))
                elif i==len(a)-1:
                    sub.append(float(a[i][:-1]))
                else:
                    sub.append(float(a[i]))
            hold.append(np.array(sub))
        check_len=True
        check_zero=True
        check_lin=False
        veco=np.array([0,0,0,0])
        new=[]
        for item in hold:
            mul=np.dot(mata,item)
            if np.array_equal(np.dot(mata,item),np.array([0,0,0,0])) and not np.array_equal(item,veco):
                new.append(item)
            elif not np.array_equal(np.dot(mata,item),np.array([0,0,0,0])):
                check_zero=False
                res['msg']='At least one of the vectors is not in the kernel'
        if (len(new)==1 or len(new)==0) and check_zero==True:
            res['msg']='The set of vectors given does not span the kernel'
        else:
            for i in range(len(new)-1):
                for j in range(i+1,len(new)):
                    if span2(new[i],new[j])==1:
                        check_lin=True
                        break
            if check_lin==True:
                res['ok']=True
            elif check_zero==True and len(new)&gt;=2:
                res['msg']='You may have linearly dependent vectors'
        return res

</script>
    <p>
      <customresponse cfn="spangrader" expect="&amp;lt;1,-1,0,1&amp;gt;;&amp;lt;-1,2,1,0&amp;gt;">
        <textline correct_answer="&amp;lt;1,-1,0,1&amp;gt;;&amp;lt;-1,2,1,0&amp;gt;"/>
      </customresponse>
    </p>
    <p>
      <solution>
        <font color="blue">Answer: </font>
        <font color="blue">
          <p>
 We row reduce [mathjaxinline]A[/mathjaxinline] and obtain </p>
          <table id="a0000000175" class="equation" width="100%" cellspacing="0" cellpadding="7" style="table-layout:auto">
            <tr>
              <td class="equation" style="width:80%;vertical-align:middle;text-align:center;border-style:hidden">[mathjax]\left[ \begin{array}{cccc} 1 &amp;  0 &amp;  1 &amp;  -1 \\ 0 &amp;  1 &amp;  -2 &amp;  1 \\ 0 &amp;  0 &amp;  0 &amp;  0 \\ 0 &amp;  0 &amp;  0 &amp;  0 \end{array} \right][/mathjax]</td>
              <td class="eqnnum" style="width:20%;vertical-align:middle;text-align:left;border-style:hidden"/>
            </tr>
          </table>
          <p>
Therefore, solutions to [mathjaxinline]Ax = \mathbf{0}[/mathjaxinline] must be of the form </p>
          <table id="a0000000176" class="equation" width="100%" cellspacing="0" cellpadding="7" style="table-layout:auto">
            <tr>
              <td class="equation" style="width:80%;vertical-align:middle;text-align:center;border-style:hidden">[mathjax]\left[ \begin{array}{c} x_1 \\ x_2 \\ x_3 \\ x_4 \end{array} \right] = \left[ \begin{array}{c} -x_3 + x_4 \\ 2x_3 - x_4 \\ x_3 \\ x_4 \end{array} \right] = x_3 \left[ \begin{array}{c} -1 \\ 2 \\ 1 \\ 0 \end{array} \right] + x_4 \left[ \begin{array}{c} 1 \\ -1 \\ 0 \\ 1 \end{array} \right],[/mathjax]</td>
              <td class="eqnnum" style="width:20%;vertical-align:middle;text-align:left;border-style:hidden"/>
            </tr>
          </table>
          <p>
 where [mathjaxinline]x_3[/mathjaxinline] and [mathjaxinline]x_4[/mathjaxinline] are free variables. Hence the list </p>
          <table id="a0000000177" class="equation" width="100%" cellspacing="0" cellpadding="7" style="table-layout:auto">
            <tr>
              <td class="equation" style="width:80%;vertical-align:middle;text-align:center;border-style:hidden">[mathjax]\left\{  \left[ \begin{array}{c} -1 \\ 2 \\ 1 \\ 0 \end{array} \right]; \left[ \begin{array}{c} 1 \\ -1 \\ 0 \\ 1 \end{array} \right] \right\}[/mathjax]</td>
              <td class="eqnnum" style="width:20%;vertical-align:middle;text-align:left;border-style:hidden"/>
            </tr>
          </table>
          <p>
 forms a basis of the kernel of [mathjaxinline]T[/mathjaxinline].  </p>
        </font>
      </solution>
    </p>
  </text>
</problem>
