<problem display_name="Closest Solution" linenum="136" tex_filename="./hw/hw4.tex" showanswer="finished" attempts="4" rerandomize="per_student">
  <text>
    <p>
Let [mathjaxinline]A = \left[\begin{array}{cc} -2 &amp;  1 \\ 1 &amp;  -1 \\ 1 &amp;  2 \end{array} \right].[/mathjaxinline] </p>
    <p>
Find the vector [mathjaxinline]x[/mathjaxinline] for which [mathjaxinline]Ax[/mathjaxinline] comes closest to [mathjaxinline]v=\left[ \begin{array}{c} 2 \\ -3 \\ 2 \end{array} \right].[/mathjaxinline] </p>
    <script linenum="2" filename="./vectorentry2.tex" type="text/python" system_path="python_lib">
def VectorEntry(expect, ans):
	import ast
	import numpy as np 
  	atol = 0.01
	list_expect = ast.literal_eval(expect)
	vec_expect = np.matrix(list_expect)
  	ret = {"ok":False}
	try:
  		# input format [[1],[2],[3]]
		list_ans = ast.literal_eval(ans)
		vec_ans = np.matrix(list_ans)
  		if vec_ans.shape != vec_expect.shape:
  			ret['msg'] = 'Wrong shape of vector!'
  		elif np.allclose(vec_ans, vec_expect,atol,1e-08):
  			ret['ok'] = True
  		else:
  		# More error message. Will improve this part
  			ret['msg'] = 'something is wrong'   			
	#except SyntaxError:
		#ret['msg'] = 'Wrong input format'
	except SyntaxError:
  		# input format &amp;lt;1,2,3&amp;gt;
		list_ans = ans.replace('&amp;lt;', '[').replace('&amp;gt;', ']')
		list_ans = ast.literal_eval(list_ans)
		vec_ans = np.transpose(np.matrix(list_ans))
  		if vec_ans.shape != vec_expect.shape:
  			ret['msg'] = 'Wrong shape of vector!'
  		elif np.allclose(vec_ans, vec_expect,0.01,1e-08):
  			ret['ok'] = True
  		else:
    		# More error message. Will improve this part
  			ret['msg'] = 'something is wrong' 
  	except:
  		ret['msg'] = 'Wrong input format'
  	return ret 
</script>
    <p>
      <customresponse cfn="VectorEntry" expect="[[-.6],[1.4]]">
        <textline correct_answer="[[-.6],[1.4]]"/>
      </customresponse>
    </p>
    <p>
      <solution>
        <font color="blue">Answer: </font>
        <font color="blue">
          <p>
 We seek the least-squared solution to [mathjaxinline]Ax = v[/mathjaxinline]. We first calculate </p>
          <table id="a0000000212" class="equation" width="100%" cellspacing="0" cellpadding="7" style="table-layout:auto">
            <tr>
              <td class="equation" style="width:80%;vertical-align:middle;text-align:center;border-style:hidden">[mathjax]A^ tA = \left[\begin{array}{cc} 6 &amp;  -1 \\ -1 &amp;  6 \end{array} \right],[/mathjax]</td>
              <td class="eqnnum" style="width:20%;vertical-align:middle;text-align:left;border-style:hidden"/>
            </tr>
          </table>
          <p>
 and [mathjaxinline]A^ tv = \left[ \begin{array}{c} -5 \\ 9 \end{array} \right].[/mathjaxinline] </p>
          <p>
The least-squared solution is then [mathjaxinline](A^ tA)^{-1}A^ tv = \left[ \begin{array}{c} -0.6 \\ 1.4 \end{array} \right].[/mathjaxinline] </p>
        </font>
      </solution>
    </p>
  </text>
</problem>
